<h1>PHY 2200 - Computational Physics</h1>
<h2>Spring 2023</h2>

## Python, Jupyter, and Markdown

Welcome to the course! The goals for today are to make sure you have the Anaconda distrbution of Python installed, to introduce you to the workflow of this class, and to give you a little bit of exposure to using Markdown to writing mathematical expressions. 

Communication is critical. When you are working on projects throughout this course, you will need to explain what you are doing beyond just commenting the code. You are telling a story, and your story will necessarily involve computation. Markdown allows you to use $\LaTeX$ commands to typeset equations, special characters, and other mathematical statements. You might not be familiar with $\LaTeX$, but a great way to get practice is to sign up for a free account with [Overleaf](https://www.overleaf.com/). We won't need to generate PDF documents or use much of the fancy typesetting that the full package uses. But we will make use of basic commands and operations to make equations.

This might seem an odd way to start off a course on computational physics. It is. We're getting to know each other, and I already know this is a diverse group. Some of you have little programming experience, and some of you have never taken physics. Think of the first few days as a trial period where we will build up some expertise in a set of common skills.

The cell containing this and the previous three paragraphs is a Markdown cell. If you double click, you can see the source. An oversimplified (but <i>useful</i>) way of thinking about Markdown is as simpler (but more limited) version of HTML in which you can write equations that look good. Today, we'll go through some examples of things you can do. In the following cell, you can see several equations from physics and higher mathematics. Double click to see how these are created:

# Here is a list of equations:

$$\frac{d{\bf p}}{dt} = {\bf F}_{ext}$$

$$\nabla \cdot {\bf E} = \frac{\rho({\bf r})}{\epsilon_{0}}$$

$$\nabla \times {\bf E} = -\frac{\partial {\bf B}}{\partial t}$$

$$\gamma = \frac{1}{\sqrt{1-\frac{{\bf v}\cdot {\bf v}}{c^{2}}}}$$

$$\int_{-\infty}^{\infty}e^{-\alpha x^{2}}dx = \sqrt{\frac{\pi}{\alpha}}$$

$$\langle \cos^{2}\theta\rangle = \frac{\int_{0}^{2\pi}\cos^{2}\theta d\theta}{\int_{0}^{2\pi}d\theta} = \frac{1}{2}$$

$$E^{2} = \sqrt{m^{2}c^{4} + p^{2}c^{2}}$$

$$\frac{\partial f}{\partial t} = \nabla^{2}f$$

$$\sigma^{y} = \left(\begin{array}{cc} 0 & -i \\ i & 0 \end{array}\right)$$

$${\bf I} = \left(\begin{array}{ccc} I_{xx} & I_{xy} & I_{xz}\\ I_{yx} & I_{yy} & I_{yz}\\ I_{zx} & I_{zy} & I_{zz}\end{array}\right)$$

$$i\hbar\frac{d|\psi\rangle}{dt} = \hat{H}|\psi\rangle$$


$$\sum_{n=1}^{\infty}\frac{1}{n^{2}} = \frac{1}{1^{2}} + \frac{1}{2^{2}} + \frac{1}{3^{2}} + \cdots  = \frac{\pi^{2}}{3!}$$

$$\Delta p \Delta x \geq \frac{\hbar}{2}$$

$$f(x+ \delta x) \simeq f(x) + \delta x f'(x) + \frac{\delta x^{2}}{2}f''(x)$$

$$\Gamma(z)\Gamma(1-z) = \frac{\pi}{\sin(\pi z)}$$

$$\eta_{\mu\nu} = \left(\begin{array}{cccc} 1 & 0 &0 & 0\\ 0 & -1 & 0 & 0 \\ 0 & 0 & -1 & 0\\ 0 & 0 & 0 & -1\end{array}\right)$$


$$\mathcal{L} = -\frac{1}{4}F_{\mu\nu}F^{\mu\nu}$$


$$\LaTeX\mbox{ is life}$$

Got it? Great.

But really, this should look overwhelming the first time you see it. Google is your friend. For example, [here](https://artofproblemsolving.com/wiki/index.php/LaTeX:Symbols) is one of the first things that came up when I searched for a quick reference sheet. It gives you probably every symbol and structure you'd need for most applications in this course. Depending on what you're trying to do, you might not find the exact answer here. Look a little harder, or ask me for help.

For your first task, take a look at the equations below. I pasted <i>images</i> of these, so you can't see the $\LaTeX$ source when you double click. In the cell below this, figure out how to write each equation using $\LaTeX$ in a Markdown cell.

<b>Important: </b> By default, each new cell is initialized as a `code` cell, meaning Python commands are expected. You get an error trying to type out equations. See the `code` option on the pulldown menu at the top? You can change this to `Markdown` to let Jupyter know to expect Markdown input that it will interpret and typeset quite nicely for you. Don't be hard on yourself. You <i>will</i> forget to change these cell types from time to time. It will not break Python.

<center>
<img src="eq1.png" width=300>
<img src="eq2.png" width=500>
</center>

## Solution:

$$(\nabla \times A)_i = \epsilon_{ijk}\frac{\partial A_k}{\partial x_j}$$
$$\lim _{x \to \infty} \frac{\sin x}{x} = \frac{\pi}{2}$$
$$i \bar{h} \gamma^\mu \partial_\mu \psi (x) - m c \psi (x) = 0$$
$$\sigma^x = \begin{pmatrix} 0 & 1 \\ 1 & 0 \end{matrix}$$
$$\nabla \cdot \bf B = 0$$
$$\nabla \times \bf B = \mu_0 (\bf J + \epsilon_0 \frac{\partial \bf E}{\partial t}$$
$$I_{jk} = \sum m_n [ r^2 \delta_{jk} - x_j x_k ]$$
$$F_{\mu v} = \partial_\mu A_v - \partial_v A_\mu$$
$$\aleph_\omega = sup { \aleph_n : n \in \omega }$$
$$S = k_B ln \Omega$$
$$e^10 = cos \theta + i sin \theta$$
$$R_{\mu v} - \frac{1}{2} Rg_{\mu v} + \wedge g_{\mu v} = \frac{8 \pi G}{c^4} T_{\mu v}$$
$$\Sigma \quad \sum \frac{1}{n^3} = \frac{1}{1^3} + \frac{1}{2^3} + \frac{1}{3^3} + ... + \varsigma (3)$$
$$\gamma = \lim _{n \to \infty} [ \Sigma \quad \sum \frac{1}{k} - ln \emph{n} ] = 0.577$$
$$e^x = 1 + x + \frac{x^2}{2!} + \frac{x^3}{3!} + \Theta (x^4)$$
$$\bar{Z} [ J ] = \int D \phi exp [ \int d^4 x ( \frac{1}{2} ( \partial_\mu \phi \partial^u \phi - m^2 \phi^2 ) + J \phi) ] $$
$$g = \begin{pmatrix} 1 & 0 & 0 & 0\\ 0 & -1 & 0 & 0 \\ 0 & 0 & -r^2 & 0\\ 0 & 0 & 0 & -r^2 sin^2 \theta \end{pmatrix}$$
$$\mid \Psi ( t ) \rangle = \Sigma \langle \phi_n \mid \Psi ( 0 ) \rangle e^{-i E_n t l h} \mid \Phi_n \rangle$$

Another feature we'll use in this class is the `nbgrader` package for evaluating notebooks. This doesn't require <i>you</i> to install anything specific. But your answers (solutions to problems, segments of code, etc.) will need to be placed in particular places. Most often, you'll be completing <i>functions</i> which return some result. To explore the workflow, let's look at some of the classic "math problems for children" from mathematical physicist Vladimir Arnold (available [here](https://www.imaginary.org/sites/default/files/taskbook_arnold_en_0.pdf) for your amusement).

<h3>1. A brick weighs one pound and half the brick. How many pounds does the brick weigh?</h3>

Just below the tag `# YOUR CODE HERE`, write your answer in the form `return x` where `x` is replaced by what you think the answer is.

<!-- grade_id: cell-851c4a976bc96ee9, points: 0, grade: false, solution: true -->
```python

def answer():
    return 2
    raise NotImplementedError()
```

The first trick is making sure that works without error. That does <i>not</i> necessarily mean your answer is correct. Below you can see the actual testing feature. For some questions, you'll be able to check if your solution passes the test. There will also (sometimes) be hidden tests where you will not be able to see how your solution is being checked. For these silly math problems, you can at least see that some sort of test is being performed.

If you can execute the following block with no error, you got the answer right!

<!-- grade_id: cell-9141513aba05fe6a, points: 5, grade: true, solution: false -->
```python
'''case 1'''
import math
import numpy as np
f = open('tests.csv','r')
A = np.genfromtxt(f, delimiter=",")
A = A[1]
f.close()
assert( math.isclose(answer(),A))
```

Now try your luck at as many of these as you have time to attempt. Embrace the intellectual struggle in solving problems supposedly designed for children!

<h3>2. A bottle with a cork costs 10 kopecks, while the bottle itself is 9 kopecks more expensive than the cork. How much does the bottle without the cork cost?</h3>

<!-- grade_id: cell-Q2, points: 0, grade: false, solution: true -->
```python
def answer():
    return 9.5
    raise NotImplementedError()
```

<!-- grade_id: cell-Q2eval, points: 5, grade: true, solution: false -->
```python
'''case 1'''
import math
import numpy as np
f = open('tests.csv','r')
A = np.genfromtxt(f, delimiter=",")
A = A[2]
f.close()
assert( math.isclose(answer(),A))
```

<h3>3. Two old ladies left from 
A to B and from B to A at dawn heading towards one another (along the same road). They met at noon, but did not stop, and each of them carried on walking with the same speed. The ﬁrst lady came (to B) at 4 p.m., and the second (to A) at 9 p.m. How many hours before noon was that day?</h3>

<!-- grade_id: cell-9eff9c573574adc1, points: 0, grade: false, solution: true -->
```python
def answer():
    return 6
    raise NotImplementedError()
```

<!-- grade_id: cell-08313e36d2e4388a, points: 0, grade: true, solution: false -->
```python
'''case 1'''
import math
import numpy as np
f = open('tests.csv','r')
A = np.genfromtxt(f, delimiter=",")
A = A[3]
f.close()
assert( math.isclose(answer(),A))
```

<h3>4. Vasya has 2 sisters more than he has brothers. How many daughters more than sons do Vasya’s parents have?

<!-- grade_id: cell-25e64992bee3358f, points: 0, grade: false, solution: true -->
```python
def answer():
    return 1
    raise NotImplementedError()
```

<!-- grade_id: cell-f553b231a441d3bf, points: 0, grade: true, solution: false -->
```python
'''case 1'''
import math
import numpy as np
f = open('tests.csv','r')
A = np.genfromtxt(f, delimiter=",")
A = A[4]
f.close()
assert( math.isclose(answer(),A))
```

<h3>5. During the daytime a snail climbs 3 cm up a post, and during the night, falling asleep, accidentally goes down by 2 cm. The post is 10 m high, and a delicious (for the snail) sweet is on its top. In how many days will the snail get the sweet?</h3>

<!-- grade_id: cell-39b4949d36a3529a, points: 0, grade: false, solution: true -->
```python
def answer():
    return 998
    raise NotImplementedError()
```

<!-- grade_id: cell-e6e4c6536f82a604, points: 0, grade: true, solution: false -->
```python
'''case 1'''
import math
import numpy as np
f = open('tests.csv','r')
A = np.genfromtxt(f, delimiter=",")
A = A[5]
f.close()
assert( math.isclose(answer(),A))
```

<h3>6. Two volumes of Pushkin, the ﬁrst and the second, are side-by-side on a bookshelf. The pages of each volume are 2 cm thick, and the cover - front and back each - is 2
mm. A bookworm has gnawed through (perpendicular to the pages) from the ﬁrst page of volume 1 to the last page of volume 2. How long (in cm) is the bookworm’s track?</h3>

<!-- grade_id: cell-fc7141694c25d85f, points: 0, grade: false, solution: true -->
```python
def answer():
    return .4
    raise NotImplementedError()
```

<!-- grade_id: cell-5285d3359545281e, points: 0, grade: true, solution: false -->
```python
'''case 1'''
import math
import numpy as np
f = open('tests.csv','r')
A = np.genfromtxt(f, delimiter=",")
A = A[6]
f.close()
assert( math.isclose(answer(),A))
```

<h3>7. The distance between cities A and B is 40km. Two cyclists leave respectively from A and B simultaneously towards one another, one with speed 10km/h and the other with speed 15km/h. A ﬂy ﬂies out with the ﬁrst cyclist from A with the speed of 100km/h, reaches the second, touches his forehead and ﬂies back to the ﬁrst, touches his forehead,returns to the second,and so on until the cyclists’ foreheads collide and squash the ﬂy. How many kilometers altogether has the ﬂy ﬂown?</h3>

<!-- grade_id: cell-203420f0e8367ea8, points: 0, grade: false, solution: true -->
```python
def answer():
    return 160
    raise NotImplementedError()
```

<!-- grade_id: cell-c23a4435e17bf5c1, points: 0, grade: true, solution: false -->
```python
'''case 1'''
import math
import numpy as np
f = open('tests.csv','r')
A = np.genfromtxt(f, delimiter=",")
A = A[7]
f.close()
assert( math.isclose(answer(),A))
```

<h3>8. A caterpillar wants to slither from a corner of a cubic room (the left on the ﬂoor) to the opposite one (the right on the ceiling). Find the shortest route (in meters) for such a journey along the walls of the room. Each side of the room is 2 m.</h3>

<!-- grade_id: cell-2aeeda30a83fee84, points: 0, grade: false, solution: true -->
```python
def answer():
    return 4.472135955
    raise NotImplementedError()
```

<!-- grade_id: cell-f0d7dd09c276c307, points: 0, grade: true, solution: false -->
```python
'''case 1'''
import math
import numpy as np
f = open('tests.csv','r')
A = np.genfromtxt(f, delimiter=",")
A = A[8]
f.close()
assert( math.isclose(answer(),A,abs_tol=1e-3))
```

<h3>9. There are ﬁve heads and fourteen legs in a family. How many people and how many dogs are in the family?</h3>

<!-- grade_id: cell-82159c113d21b0e9, points: 0, grade: false, solution: true -->
```python
def answerpeople():
    return 3
    raise NotImplementedError()
    
def answerdogs():
    return 2
    raise NotImplementedError()
```

<!-- grade_id: cell-c793330e8781b64e, points: 0, grade: true, solution: false -->
```python
'''case 1'''
import math
import numpy as np
f = open('tests.csv','r')
A = np.genfromtxt(f, delimiter=",")
A = A[9]
f.close()
assert( math.isclose(answerpeople(),A,abs_tol=1e-3))
```

<!-- grade_id: cell-3cff1cabdcc45b92, points: 0, grade: true, solution: false -->
```python
'''case 1'''
import math
import numpy as np
f = open('tests.csv','r')
A = np.genfromtxt(f, delimiter=",")
A = A[10]
f.close()
assert( math.isclose(answerdogs(),A,abs_tol=1e-3))
```

<h3>10. Find the limit of the following infinitely continued fraction:
</h3>

$$x = 1+\frac{1}{2+\frac{1}{1+\frac{1}{2+\frac{1}{1+\frac{1}{2+\cdots}}}}}$$

<i> Hint: </i> use self-similarity

<!-- grade_id: cell-249932f88c13e0d2, points: 0, grade: false, solution: true -->
```python
def answer():
    return 1.366025404
    raise NotImplementedError()
```

<!-- grade_id: cell-890fadb0037ca5d7, points: 0, grade: true, solution: false -->
```python
'''case 1'''
import math
import numpy as np
f = open('tests.csv','r')
A = np.genfromtxt(f, delimiter=",")
A = A[11]
f.close()
assert( math.isclose(answer(),A,abs_tol=1e-3))
```

```python

```

```python

```